The integrated WF-Haldane (WFH) model of genetic drift resolving the many paradoxes of molecular evolution

Author response:

The following is the authors’ response to the original reviews.

eLife Assessment

This study presents a useful modification of a standard model of genetic drift by incorporating variance in offspring numbers, claiming to address several paradoxes in molecular evolution. It is unfortunate that the study fails to engage prior literature that has extensively examined the impact of variance in offspring number, implying that some of the paradoxes presented might be resolved within existing frameworks.

The prior literature the reviewers referred to are all "modified WF models". In the original submission, we lumped the standard and modified WF models together as the "generalized WF models". As the lumping causes confusions, their distinctions are now made clear. That said, the Haldane model in our proposal is not a modification of the standard WF model because, conceptually, the two models are very different. WF is based on sampling whereas the Haldane model is based on gene transmission.

While the "modified WF models" often incorporate V(K) [variance in progeny number], the modification is still based on the WF model of population sampling. The modification is mathematically feasible but biologically untenable, as explained explicitly in the revised text. Most important, all four paradoxes are as incompatible with the modified WF models as with the standard model. Note that the Haldane model does not have the sampling step, which is absorbed into the V(K) term. In the integrated WF-Haldane model, these paradoxes are resolved (see the new sections of Discussion, quoted below).

If readers do not have time to ponder on all four paradoxes, they may simply read the first one, as follows. When the population size (N) is growing exponentially, such as in a bacteria culture, drift is nearly absent when N is small and becomes stronger as N increases, especially when approaching the carrying capacity. Such common observations are exactly opposite of the WF model's central prediction. Any model based on sampling cannot escape the constraint of "greater drift, smaller N".

Revision - The following text is a reproduction of the last 7 paragraphs of Discussion.

“The standard WF model has been extended in several directions (overlapping generations, multiple alleles, ploidy, etc.). The modification most relevant to our studies here is the introduction of V(K) into the model, thus permitting V(K) ≠ E(K). While the modifications are mathematically valid, they are often biologically untenable. Kimura and Crow (1963) may be the first to offer a biological mechanism for V(K) ≠ E(K), effectively imposing the Haldane model on the WF model. Other models (Kimura and Crow 1963; Lynch, et al. 1995; Sjodin, et al. 2005; Der, et al. 2011; Cannings 2016) indeed model mathematically the imposition of the branching process on the population, followed by the WF sampling. The constructions of such models are biologically dubious but, more importantly, still unable to resolve the paradoxes. It would seem more logical to use the Haldane model in the first place by having two parameters, E(K) and V(K).

Even if we permit V(K) ≠ E(K) under the WF sampling, the models would face other difficulties. For example, a field biologist needs to delineate a Mendelian population and determine its size, N or Ne. In all WF models, one cannot know what the actual population being studied is. Is it the fly population in an orchard being sampled, in the geographical region, or in the entire species range? It is unsatisfactory when a population biologist cannot identify the population being studied. The Haldane model is an individual-output model (Chen, et al. 2017), which does not require the delineation of a Mendelian population.

We shall now review the paradoxes specifically in relation to the modified WF models, starting with the multi-copy gene systems such as viruses and rRNA genes covered in the companion study (Wang, et al. 2024). These systems evolve both within and between hosts. Given the small number of virions transmitted between hosts, drift is strong in both stages as shown by the Haldane model (Ruan, Luo, et al. 2021; Ruan, Wen, et al. 2021; Hou, et al. 2023). Therefore, it does not seem possible to have a single effective population size in the WF models to account for the genetic drift in two stages. The inability to deal with multi-copy gene systems may explain the difficulties in accounting for the SARS-CoV-2 evolution (Deng, et al. 2022; Pan, Liu, et al. 2022; Ruan, Wen, et al. 2022; Hou, et al. 2023; Ruan, et al. 2023).

We now discuss the first paradox of this study, which is about the regulation of N. In the general WF models, N is imposed from outside of the model, rather than self-generating within the model. When N is increasing exponentially as in bacterial or yeast cultures, there is almost no drift when N is very low and drift becomes intense as N grows to near the carrying capacity. As far as we know, no modifications of the WF model can account for this phenomenon that is opposite of its central tenet. In the general WF models, N is really the carrying capacity, not population size.

The second paradox of sex chromosomes is rooted in V(K) ≠ E(K). As E(K) is the same between sexes but V(K) is different, clearly V(K) = E(K) would not be feasible. The mathematical solution of defining separate Ne's for males and females (Kimura and Crow 1963; Lynch, et al. 1995; Sjodin, et al. 2005; Der, et al. 2011; Cannings 2016) unfortunately obscures the interesting biology. As shown in Wang et al. (2024; MBE), the kurtosis of the distribution of K indicates the presence of super-breeder males. While the Haldane model can incorporate the kurtosis, the modified WF models are able to absorb only up to the variance term, i.e., the second moment of the distribution. The third paradox of genetic drift is manifested in the fixation probability of an advantageous mutation, 2_s_/V(K). As explained above, the fixation probability is determined by the probability of reaching a low threshold that is independent of N itself. Hence, the key parameter of drift in the WF model, N (or Ne), is missing. This paradox supports the assertion that genetic drift is fundamentally about V(K) with N being a scaling factor.

As the domain of evolutionary biology expands, many new systems do not fit into the WF models, resulting in the lack of a genetic drift component in their evolutionary trajectories. Multi-copy gene systems are obvious examples. Others include domestications of animals and plants that are processes of rapid evolution (Diamond 2002; Larson and Fuller 2014; Purugganan 2019; Chen, Yang, et al. 2022; Pan, Zhang, et al. 2022; Wang, et al. 2022). Due to the very large V(K) in domestication, drift must have played a large role. Somatic cell evolution is another example with “undefinable” genetic drift (Wu, et al. 2016; Chen, et al. 2017; Chen, et al. 2019; Ruan, et al. 2020; Chen, Wu, et al. 2022). The Haldane (or WFH) model, as an "individual output" model, can handle these general cases of genetic drift.

The Haldane model and the WF model are fundamentally different approaches to random forces of evolution. While the WF models encounter many biological contradictions, they have provided approximate mathematical solutions to more realistic scenarios. In systems such as in viral evolution (Ruan, Hou, et al. 2022; Hou, et al. 2023) or somatic cell evolution (Chen, Wu, et al. 2022; Zhai, et al. 2022) whereby the WF solution is absent, further development of the WFH model will be necessary.”

In addition, while the modified model yields intriguing theoretical predictions, the simulations and empirical analyses are incomplete to support the authors' claims.

This point is addressed in the responses to reviewers' comments. Since they are quite technical, they do not fit in the overview here.

Public Reviews:

Reviewer #1 (Public Review):

Summary:

The authors present a theoretical treatment of what they term the "Wright-Fisher-Haldane" model, a claimed modification of the standard model of genetic drift that accounts for variability in offspring number, and argue that it resolves a number of paradoxes in molecular evolution. Ultimately, I found this manuscript quite strange.

The notion of effective population size as inversely related to the variance in offspring number is well known in the literature, and not exclusive to Haldane's branching process treatment. However, I found the authors' point about variance in offspring changing over the course of, e.g. exponential growth fairly interesting, and I'm not sure I'd seen that pointed out before.

Weaknesses:

I have several outstanding issues. First of all, the authors really do not engage with the literature regarding different notions of an effective population. Most strikingly, the authors don't talk about Cannings models at all, which are a broad class of models with non-Poisson offspring distributions that nonetheless converge to the standard Wright-Fisher diffusion under many circumstances, and to "jumpy" diffusions/coalescents otherwise (see e.g. Mohle 1998, Sagitov (2003), Der et al (2011), etc.). Moreover, there is extensive literature on effective population sizes in populations whose sizes vary with time, such as Sano et al (2004) and Sjodin et al (2005).

Of course in many cases here the discussion is under neutrality, but it seems like the authors really need to engage with this literature more.

The reviewer's summary and weakness statement reflects the general criticism summarized by the editors. The reply and revision to these criticisms have been presented in the long reply to elife assessment above.

We hence re-emphasize only the key points here.

(1) The literature that the reviewers fault us for not citing is about the modifications of the standard WF model. We now cite them as well as a few others in that vein. However, the WF-Haldane model we propose is conceptually very different from the modified WF models. This WFH model is in essence the Haldane model which may use the results of the WF models as the starting point to find the exact solutions.

(2) The check of the power of the modified WF models is whether they can resolve the paradoxes. None of them can. The arguments apply to neutral cases as well as selection effects. Hence, our central point is that the modifications of the standard WF model [e.g., by incorporating V(K)] do not help the WF model in resolving the paradoxes. Besides, the incorporation of V(K) is mathematically feasible but biologically untenable as presented in the new sections of Discussion.

Nonetheless, I don't think the authors' modeling, simulations, or empirical data analysis are sufficient to justify their claims.

The most interesting part of the manuscript, I think, is the discussion of the Density Dependent Haldane model (DDH). However, I feel like I did not fully understand some of the derivation presented in this section, …… - this is the whole notion of exchangeability, also neglected in this manuscript). As such, I don't believe that their analysis of the empirical data supports their claim. [Since the comments above are highly technical and fairly long, they are not copied verbatim.]

We thank this reviewer for the detailed comments with respect to the potential confusion in the discussion of the Density Dependent Haldane (DDH) model.

First, the reviewer appears to ask how Eqs (5-6) are derived. We should clarify that both Eq (5) and (6) are assumptions rather than derived results. Both equations are assumptions based on population ecology. Eq (7) is then derived by substituting the assumptions in Eq (5) and (6) into Eq (3).

The definition in Equation (5) allows the growth rate of the population size to be dependent on N itself, such that growth rate E(K) (average offspring number per generation) is greater than 1 when N < Ck and less than 1 when N > Ck. The parameter z is introduced to adjust the sensitivity of E(K) to changes in population size (as shown in Fig. 3a).

Second, we appreciate the comments regarding the use of individual-based simulations and the apparent lack of interaction between individuals. In our simulations, there is indeed an interaction among individuals, which is represented by Eq (5). This equation reflects how the competition between two alleles affects the expected growth rate 𝐸(𝐾), which decreases as the population size increases. Furthermore, once 𝐸(𝐾) for the entire population is determined, the offspring numbers of the alleles are independent.

We believe that the primary purpose of our simulations was not clearly stated. This lack of clarity may be the root of the criticisms. We now note that the simulations are aimed at testing the accuracy of Equation (10).

Note that Eq. (10) is a textbook result and quite important in our study. This equation shows that the strength of genetic drift, as given by Pf (the fixation probability of an advantageous mutation), is not a function of N at all. This approximate solution has been obtained using the WF model by Kimura. The Haldane model solution that can explain Paradox 1 is based on Equation (7) as shown below

Since the fixation probability of Equation (10) cannot be easily obtained using Eq. (7), we conducted simulations to confirm the accuracy of Eq. (10) when applied to the Haldane model.

We have revised the relevant sections of the manuscript to clarify these points and to better distinguish between assumptions and results.

Revision - Details of the DDH model are given in the Supplementary Information. A synopsis is given here: We consider a non-overlapping haploid population with two neutral alleles. The population size at time t is Nt. We assume that expected growth rate E(K) is greater than 1 when N < Ck and less than 1 when N > Ck, as defined by Eq. (5) below:

The slope of E(K) vs. N (i.e., the sensitive of growth rate to changes in population size), as shown in Fig 3a, depends on z. To determine the variance V(K), we assume that K follows the negative binomial distribution whereby parents would suffer reproduction-arresting injury with a probability of pt at each birthing (Supplementary Information). Accordingly, V(K) can then be expressed as

By Eq. (6), the ratio of V(K)/E(K) could be constant, decrease or increase with the increase of population size. With E(K) and V(K) defined, we could obtain the effective population size by substituting Eq. (5) and Eq. (6) into Eq. (3).

Eq. (7) presents the relationship between effective population size (Ne) and the population size (N) as shown in Fig. 3. The density-dependent E(K) could regulate N with different strength (Fig. 3a). The steeper the slope in Fig. 3a, the stronger the regulation.

Simulation of genetic drift in the Haldane model and the Wright-Fisher (WF) model. In both models, interactions between individuals are implicitly included through the dependency of the average number of offspring on population size, as defined by Eq. (5). This dependency leads to the logistic population growth, reflecting the density-dependent interactions.

Thus, while I think there are some interesting ideas in this manuscript, I believe it has some fundamental issues:

first, it fails to engage thoroughly with the literature on a very important topic that has been studied extensively. Second, I do not believe their simulations are appropriate to show what they want to show. And finally, I don't think their empirical analysis shows what they want to show.

References omitted

The comments are the summary of previous ones, which have been addressed in detail in the preceding sections.

Reviewer #2 (Public Review):

Summary:

This theoretical paper examines genetic drift in scenarios deviating from the standard Wright-Fisher model. The authors discuss Haldane's branching process model, highlighting that the variance in reproductive success equates to genetic drift. By integrating the Wright-Fisher model with the Haldane model, the authors derive theoretical results that resolve paradoxes related to effective population size [Ne]

Thanks. The issue of Ne will be addressed below where the reviewer returns to this issue. The strength of the integrated WFH model is that N (or Ne) is generated by the model itself, rather than externally imposed as in WF models.

Strengths:

The most significant and compelling result from this paper is perhaps that the probability of fixing a new beneficial mutation is 2s/V(K). This is an intriguing and potentially generalizable discovery that could be applied to many different study systems.

The authors also made a lot of effort to connect theory with various real-world examples, such as genetic diversity in sex chromosomes and reproductive variance across different species.

Thanks.

Weaknesses:

One way to define effective population size is by the inverse of the coalescent rate. This is where the geometric mean of Ne comes from. If Ne is defined this way, many of the paradoxes mentioned seem to resolve naturally. If we take this approach, one could easily show that a large N population can still have a low coalescent rate depending on the reproduction model. However, the authors did not discuss Ne in light of the coalescent theory. This is surprising given that Eldon and Wakeley's 2006 paper is cited in the introduction, and the multiple mergers coalescent was introduced to explain the discrepancy between census size and effective population size, superspreaders, and reproduction variance - that said, there is no explicit discussion or introduction of the multiple mergers coalescent.

The Haldane model treats N’s very differently from the WF models. In the WF models, N’s are imposed externally (say, constant N, exponentially growing N, temporally fluctuating N’s and so on; all provided from outside of the model). Ne and coalescence are all derived from these given N’s. In order to account for the first paradox (see the next paragraph), N needs to be regulated but the WF models cannot regulate N’s. The density-dependent Haldane model that Reviewer 1 inquired above is a model that regulates N internally. It can thus account for the paradox.

Paradox 1 - When the population size (N) is growing exponentially, such as in a bacteria culture, drift is nearly absent when N is small and is much stronger as N increases, especially when approaching the carrying capacity. Such a pattern is a common observation and is exactly opposite of the WF model's central prediction. In short, a model that does not regulate N cannot explain the paradox

Ne is a fix of the WF model in order to account for the missing components of genetic drift. The paradoxes presented in this one and the companion study show that the fix is rather inadequate. In contrast, by the WFH model, N is regulated within the model itself as E(K) and V(K) are both functions of N.

The Wright-Fisher model is often treated as a special case of the Cannings 1974 model, which incorporates the variance in reproductive success. This model should be discussed. It is unclear to me whether the results here have to be explained by the newly introduced WFH model, or could have been explained by the existing Cannings model. The abstract makes it difficult to discern the main focus of the paper. It spends most of the space introducing "paradoxes".

We appreciate greatly the illuminating advice. Nevertheless, we should explain, or should have explained, more clearly that these four paradoxes presented are central to this pair of eLife papers. The WF and Haldane models are very different conceptual ideas altogether. The choice should not be based on mathematical grounds but on how they help us understand biological evolution. We are using four paradoxes to highlight the differences. We have said in the papers that the origin and evolution of COVID-19 caused a lot of confusions partly because the WF models cannot handle multi-copy gene systems, including viruses that evolve both within- and between- hosts.

The standard Wright-Fisher model makes several assumptions, including hermaphroditism, non-overlapping generations, random mating, and no selection. It will be more helpful to clarify which assumptions are being violated in each tested scenario, as V(K) is often not the only assumption being violated. For example, the logistic growth model assumes no cell death at the exponential growth phase, so it also violates the assumption about non-overlapping generations.

We appreciate the question which has two aspects. First, why do we think the WF models are insufficient? After all, for each assumption of the WF model (as given in the reviewer’s examples), there is often a solution by modifying Ne which relaxes the assumption. In this sense, there is only one grand assumption made by the WF models. That is, however complex the biology is, it is possible to find Ne that can make the WF model work. Our argument is that Ne is a cumbersome fix of the WF model and it does not work in many situations. That is how we replied about the importance of the paradoxes above. We shall again use the first paradox as an example whereby drift is stronger as N becomes larger, the fix has to make Ne negatively correlated with N. In reality, it does not appear possible to resolve this paradox. Another paradox is the evolution of multi-copy gene systems. In short, it seems clear that Ne is not a useful or usable fix.

The second aspect is that “why, among the many modifications the WF models make, do we only emphasize the inclusion of V(K)?” This is the essence of the two papers of ours. Although V(K) is a modification of the WF models, it does not enable the WF models to resolve the paradoxes. In contrast, the Haldane model has incorporate E(K) and V(K) in the model. In presenting paradox 3, it was stated that

This equation shows that the strength of genetic drift, as given by Pf (the fixation probability of an advantageous mutation), is not a function of N at all. It supports the view that the essence of genetic drift is V(K) with N as a scaling factor. Note that, if V(K) = 0, there is no genetic drift regardless of N. As V(K) is not an add-on to the Haldane model (unlike in WF models), the Haldane model can resolve the paradoxes.

The theory and data regarding sex chromosomes do not align. The fact that \hat{alpha'} can be negative does not make sense. The authors claim that a negative \hat{alpha'} is equivalent to infinity, but why is that? It is also unclear how theta is defined. It seems to me that one should take the first principle approach e.g., define theta as pairwise genetic diversity, and start with deriving the expected pair-wise coalescence time under the MMC model, rather than starting with assuming theta = 4Neu. Overall, the theory in this section is not well supported by the data, and the explanation is insufficient.

a' can be negative for the same reason that a (the male/female ratio in mutation rate) can be negative (Miyata, et al. 1987; Li, et al. 2002; Makova and Li 2002). Clearly, this has not been a problem in the large literature on a becoming negative. In fact, in many reports, a is negative, which is read as a approaching infinity. Imagine that our equation is a'^2 = 0.25, then a' can be 0.5 or -0.5, although the latter solution is not biologically meaningful.

As for theta, the reviewer asked why we do not use the pairwise genetic diversity (or theta[pi]) as the first-principle approach to estimating theta. While theta(pi) is the first estimator of theta used, the general principle is that every bin of the frequency spectrum can be used for estimating theta since the expected value is theta/i where i is the occurrence of the mutation in the sample. (If the sample size is 100, then i is between 1 and 99.) Hence, the issue is which part of the spectrum has the best statistical properties for the questions at hand. The pairwise measure is theta(pi) [which the reviewer recommends]. While theta(pi) and theta(w) are most commonly used, there are in fact numerous ways to estimate theta. ((Fu 2022) presents an excellent review.) For our purpose, we need a theta estimate least affected by selection and we choose the lowest frequency bin of the spectrum, which is theta(1) based on the singletons. Theta(1), least affected by selection, is the basis of the Fu and Li test.

Reviewer #3 (Public Review):

Summary:

Ruan and colleagues consider a branching process model (in their terminology the "Haldane model") and the most basic Wright-Fisher model. They convincingly show that offspring distributions are usually non-Poissonian (as opposed to what's assumed in the Wright-Fisher model), and can depend on short-term ecological dynamics (e.g., variance in offspring number may be smaller during exponential growth). The authors discuss branching processes and the Wright-Fisher model in the context of 3 "paradoxes": (1) how Ne depends on N might depend on population dynamics; (2) how Ne is different on the X chromosome, the Y chromosome, and the autosomes, and these differences do match the expectations base on simple counts of the number of chromosomes in the populations; (3) how genetic drift interacts with selection. The authors provide some theoretical explanations for the role of variance in the offspring distribution in each of these three paradoxes. They also perform some experiments to directly measure the variance in offspring number, as well as perform some analyses of published data.

Strengths:

(1) The theoretical results are well-described and easy to follow.

(2) The analyses of different variances in offspring number (both experimentally and analyzing public data) are convincing that non-Poissonian offspring distributions are the norm.

(3) The point that this variance can change as the population size (or population dynamics) change is also very interesting and important to keep in mind.

(4) I enjoyed the Density-Dependent Haldane model. It was a nice example of the decoupling of census size and effective size.

Thanks.

Weaknesses:

(1) I am not convinced that these types of effects cannot just be absorbed into some time-varying Ne and still be well-modeled by the Wright-Fisher process.

Please allow us to refer to, again, two of the four paradoxes. We believe that that no modification of the WF model can resolve the paradoxes.

(1) When the population size (N) is growing exponentially, such as in a bacteria culture, drift is nearly absent when N is small and is much stronger as N increases, especially when approaching the carrying capacity. Such common observations are exactly opposite of the WF model's key prediction. It is not possible for a model that does not regulate N to explain the paradox.

(2) There is no way the WF models can formulate Ne for, say viruses or ribosomal RNA genes that have two levels of populations – the within-host populations as well as the host population itself.

The fact that there are numerous Ne's suggests that Ne is a collection of cumbersome fixes of the WF model. By the WF-Haldane model, all factors are absorbed into V(K) resulting in a simpler model in the end. V(K) is often a measurable quantity. Note that, even if V(K) is incorporated into the WF model, the paradoxes remain unresolvable.

(2) Along these lines, there is well-established literature showing that a broad class of processes (a large subset of Cannings' Exchangeable Models) converge to the Wright-Fisher diffusion, even those with non-Poissonian offspring distributions (e.g., Mohle and Sagitov 2001). E.g., equation (4) in Mohle and Sagitov 2001 shows that in such cases the "coalescent Ne" should be (N-1) / Var(K), essentially matching equation (3) in the present paper.

The criticism of lack of engagement with well-established literature has been responded extensively above. Briefly, the literature is about modifications of the WF model which share the same feature of population sampling. With that feature, the paradoxes are unresolvable. For example, however Ne is defined, the fixation probability of an advantageous mutation does not depend on N or Ne. This is the third paradox of the WF models.

(3) Beyond this, I would imagine that branching processes with heavy-tailed offspring distributions could result in deviations that are not well captured by the authors' WFH model. In this case, the processes are known to converge (backward-in-time) to Lambda or Xi coalescents (e.g., Eldon and Wakely 2006 or again in Mohle and Sagitov 2001 and subsequent papers), which have well-defined forward-in-time processes.

We admire the learned understanding of the literature expressed by the review, which raise two points. First, our model may not be able to handle the heavy-tailed progeny distribution (i.e., the kurtosis of the distribution of k). Second, the Xi coalescence models (cited above) can do that. Below are our clarifications.

First, the WFH model is based on the general distribution of K, which includes flexible and realistic representations of offspring number distributions. In fact, we have used various forms of K distribution in our publications on the evolution of SARS-CoV-2 (see the Ruan et al publications in the bibliography). Power-law distribution is particularly useful as the K-distribution in viral transmission is highly kurtotic. This is reflected in the super-spreader hypothesis. In short, the branching process on which the WFH model is based in is mainly about the distribution of K. Nevertheless, the variance V(K) can often yield good approximations when the kurtosis is modest.

Second, we would like to comment on the models of Eldon and Wakely 2006. or Mohle and Sagitov 2001 and subsequent papers. These papers are based on the Moran model by considering a highly skewed distribution of offspring numbers. Fundamentally, the Moran models generally behave like WF models (standard or modified) and hence have the same problems with the paradoxes that are central to our studies. In fact, the reservations about introducing V(K) into the WF models apply as well to the Moran models. The introduction of V(K) is mathematically valid but biologically untenable. Essentially, the WF models incorporate the Haldane model as a first step in the generation transition. The introduction of V(K) into the Moran model is even less biologically sensible. Furthermore, the model allows K to take only three discrete values: 0, 2, and Nψ (see Eq. (7) in Eldon and Wakely). Their model also assumes a constant population size, which contrasts with our model's flexibility in handling varying population sizes and more complex distributions for K.

In short, the modifications of the WF (and Moran) models are unnecessarily complicated, biologically untenable but still fail to account for the paradoxes. The WFH model can rectify these problems.

(4) These results that Ne in the Wright-Fisher process might not be related to N in any straightforward (or even one-to-one) way are well-known (e.g., Neher and Hallatschek 2012; Spence, Kamm, and Song 2016; Matuszewski, Hildebrandt, Achaz, and Jensen 2018; Rice, Novembre, and Desai 2018; the work of Lounès Chikhi on how Ne can be affected by population structure; etc...)

The reviewer is correct in pointing out the inexact correlation between N and Ne. Nevertheless, it should still be true that the WF models predict qualitatively weaker drift as N increases. The first paradox is as stated:

When the population size (N) is growing exponentially, such as in a bacteria culture, drift is nearly absent when N is small and is much stronger as N increases, especially when approaching the carrying capacity. Such common observations are exactly opposite of the WF model's key prediction.

(5) I was also missing some discussion of the relationship between the branching process and the Wright-Fisher model (or more generally Cannings' Exchangeable Models) when conditioning on the total population size. In particular, if the offspring distribution is Poisson, then conditioned on the total population size, the branching process is identical to the Wright-Fisher model.

We thank the reviewer for this important comment. The main difference is that N is imposed from outside the WF models but can be generated from within the Haldane model (see the density-dependent Haldane model). In nature, N of the next generation is the sum of K’s among members of the population. It is how the Haldane model determines N(t+1) from N(t). In the WF models, N is imposed from outside the model and, hence the given N determines the distribution of K. For this reason, N regulation is not possible in the WF models, thus resulting in the paradoxes.

(6) In the discussion, it is claimed that the last glacial maximum could have caused the bottleneck observed in human populations currently residing outside of Africa. Compelling evidence has been amassed that this bottleneck is due to serial founder events associated with the out-of-Africa migration (see e.g., Henn, Cavalli-Sforza, and Feldman 2012 for an older review - subsequent work has only strengthened this view). For me, a more compelling example of changes in carrying capacity would be the advent of agriculture ~11kya and other more recent technological advances.

We thank the reviewer and have used this more convincing case as suggested by the reviewer.

Recommendations for the authors:

General replies - We thank the editors and reviewers again. The points below are re-iterations of the comments received above and have since been replied in detail. Specific instructions about wording and notations have also been rectified. Again, we are grateful for the inputs from which we learned a great deal.

Reviewing Editor Comments:

The reviewers recognize the value of this model and some of the findings, particularly results from the density-dependent Haldane model. However, they expressed considerable concerns with the model and overall framing of this manuscript.

First, all reviewers pointed out that the manuscript does not sufficiently engage with the extensive literature on various models of effective population size and genetic drift, notably lacking discussion on Cannings models and related works.

We have addressed this issue in the beginning of Introduction and Discussion, pointing to the long section in the new second half of Discussion. The essence is that the literature is all about the modified WF models. The WF-Haldane model is conceptually and operationally distinct from the WF models, either standard or modified ones,

Second, there is a disproportionate discussion on the paradoxes, yet some of the paradoxes might already be resolved within current theoretical frameworks. All three reviewers found the modeling and simulation of the yeast growth experiment hard to follow or lacking justification for certain choices. The analysis approach of sex chromosomes is also questioned.

This criticism is addressed together with the next one as they make the same point.

The reviewers recommend a more thorough review of relevant prior literature to better contextualize their findings. The authors need to clarify and/or modify their derivations and simulations of the yeast growth experiment to address the identified caveats and ensure robustness. Additionally, the empirical analysis of the sex chromosome should be revisited, considering alternative scenarios rather than relying solely on the MSE, which only provides a superficial solution. Furthermore, the manuscript's overall framing should be adjusted to emphasize the conclusions drawn from the WFH model, rather than focusing on the "unresolved paradoxes", as some of these may be more readily explained by existing frameworks. Please see the reviewers' overall assessment and specific comments.

Many thanks. We have carefully reframed and presented the WF-Haldane model to make it clear and logically consistent. Whether a new model (i.e., the WF-Haldane model) deserves to be introduced depends on whether it makes any contribution for understanding nature. That is why we emphasize the four paradoxes.

A most important disagreement between the reviewers and the authors is about the nature of the paradoxes. While the reviewers suggest that they "may" be resolvable by the conventional WF model (standard or modified), they did not offer the possible resolutions. To use the analogy in our provisional response: the WF vs. Haldane models are compared to gas cars vs electric vehicles. We can say confidently that the internal combustion engine cannot resolve the conflicting demands of transportation and zero emission. Its design has limited its capability.

Reviewer #2 (Recommendations For The Authors):

Many thanks. We have incorporated all these suggestions. When the incorporation is not straightforward, we have carefully revised the text to minimize mis-communications.

In the introduction -- "Genetic drift is simply V(K)" -- this is a very strong statement. You can say it is inversely proportional to V(K), but drift is often defined based on changes in allele frequency.

We change the word “simply” to “essentially”. This wording is supported by the fixation probability of advantageous mutations, 2s/(V(k). We have shown in the text that N does not matter here because the fixation is nearly deterministic when the copy number reaches, say, 100, regardless of whether N is 10^4 or 10^8,

Page 3 line 86. "sexes is a sufficient explanation."--> "sex could be a sufficient explanation"

The strongest line of new results is about 2s/V(K). Perhaps, the paper could put more emphasis on this part and demonstrate the generality of this result with a different example.

The math notations in the supplement are not intuitive. e.g., using i_k and j_k as probabilities. I also recommend using E[X] and V[X]for expectation and variance rather than \italic{E(X)} to improve the readability of many equations.

Thank you for your careful reading. Regarding the use of i_k and j_k as probabilities, we initially considered using 𝑝 or 𝑞 to represent probabilities. However, since 𝑝 and 𝑞 are already used in the main text, we opted for 𝑖 and 𝑗 to avoid potential confusion potential confusion. As for your recommendation to use

E[X] and V[X] for expectation and variance, we would like to clarify that we follow the standard practice of italicizing these symbols to represent variables.

Eq A6, A7, While I manage to follow, P_{10}(t) and P_{10} are not defined anywhere in the text.

Supplement page 7, the term "probability of fixation" is confusing in a branching model.

Thank you for your observation. We have carefully revised the supplement to provide clarity on these points.

Revision - In population genetics, the fixation of M allele means that the population consist entirely of the M allele, with no W alleles remaining. We define the fixation probability of M allele by generation t as follows:

Given that M and W allele reproduce independently, this can be factored as:

As t approaches infinity, the ultimate fixation probability of M allele can be derived as follows:

E.q. A 28. It is unclear eq. A.1 could be used here directly. Some justification would be nice.

We appreciate your careful review, and we will ensure this connection between the two equations is made clearer in the supplement.

Revision - Note we would like to clarify that Eq. (A1) and Eq. (A28) are essentially the same, with the only difference being the subscript 𝑡, which indicates the time dependence in the dynamic process.

Supplement page 17. "the biological meaning of negative..". There is no clear justification for this claim. As a reader, I don't have any intuition as to why that is the case.

Thank you for raising this concern. We have addressed this issue earlier.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Genetic drift, the random changes in frequencies of neutral variants, is the fundamental force of molecular evolution. Under-estimation of genetic drift is a major cause of mis-conclusions on evolution. However, the standard Wright-Fisher (WF) model of random sampling in population of size N only partially defines genetic drift with 1/N or 1/Ne (Ne being a function of varying N’s). In parallel, JBS Haldane proposed the branching process for genetic drift (Haldane 1927), whereby each gene copy is transmitted to K descendants with the mean and variance of E(K) and V(K). Genetic drift is simply V(K), i.e., the variance in transmission success, or V(K)/N when averaged over the population. Under the WF model, many paradoxes have emerged: i) Most curiously, genetic drift may often become stronger as N becomes larger at the ecological time scale, opposite to the model; ii) Sex chromosomes experience drift differently even with the same normalized N; iii) Genetic drift operates on advantageous mutations depends on V(K) but not on N. iv) Irresolution and paradoxes emerge in multi-copy gene systems, including diploidy, viruses and others, whereby evolution progresses both within and between individuals. We now show that the integration of the WF and Haldane (i.e., WFH) model can resolve these paradoxes. Most important, the WFH model can fully define genetic drift in molecular evolution.

Revised version: Genetic drift, the random changes in frequencies of neutral variants, is the fundamental force of molecular evolution. Under-estimation of genetic drift is a major cause of mis-conclusions on evolution. However, the standard Wright-Fisher (WF) model of random sampling in population of size N only partially defines genetic drift with 1/N or 1/Ne (Ne being a function of varying N’s). In parallel, JBS Haldane proposed the branching process for genetic drift (Haldane 1927), whereby each gene copy is transmitted to K descendants with the mean and variance of E(K) and V(K). Genetic drift in its essence is V(K), i.e., the variance in transmission success, or V(K)/N when averaged over the population. Under the general WF models (i.e., the standard model as well as the many modifications), paradoxes have emerged. They include: i) Most curiously, genetic drift may often become stronger as N becomes larger at the ecological time scale, opposite to the model; ii) Sex chromosomes experience drift differently even with the same normalized N; iii) Genetic drift operates on advantageous mutations depends on V(K) but not on N; iv) Irresolution and paradoxes emerge in multi-copy gene systems, including diploidy, viruses and others, whereby evolution progresses both within and between individuals (Wang, et al. 2024). We now show that the integration of the WF and Haldane (i.e., WFH) model can resolve these paradoxes. Most important, the WFH model can fully define genetic drift in molecular evolution.

Introduction

Earlier version: Genetic drift, or neutral drift, is defined as random changes in frequencies of neutral variants (Crow and Kimura 1970; Hartl and Clark 1997; Lynch, et al. 2016). Much, or even most, of DNA sequence evolution is governed by the random drift of neutral mutations. If the strength of genetic drift is under-estimated, random changes that are missed in the analysis would result in the over-estimation of other evolutionary forces of greater biological interest, including selection, mutation, migration, meiotic drive and so on. In particular, the conclusion of pervasive selection at the molecular level has been suspected to be due to the failure to fully account for genetic drift (Crow and Kimura 1970; Fu 1997; Li 1997; Charlesworth 2009; Lynch, et al. 2016; Chen, Yang, et al. 2022).

Revised version: Genetic drift is defined as the random changes in frequency of genetic variants (Crow and Kimura 1970; Hartl and Clark 1997; Lynch, et al. 2016). Much, or even most, of DNA sequence evolution is governed by the random drift of neutral mutations. Nevertheless, even non-neutral variants are subjected to genetic drift as advantageous mutations are not deterministically fixed. If the strength of genetic drift is under-estimated, random changes that are missed in the analysis would result in the over-estimation of other evolutionary forces of greater biological interest, including selection, mutation, migration, meiotic drive and so on. In particular, the conclusion of pervasive selection at the molecular level has been suspected to be due to the failure to fully account for genetic drift (Crow and Kimura 1970; Fu 1997; Li 1997; Charlesworth 2009; Lynch, et al. 2016; Chen, Yang, et al. 2022).

Earlier version: Hagedoorn and Hagedoorn (1921) may be the first to suggest that random forces (accidents, for example) could impact long-term evolution. In the same decade, models of genetic drift were developed, obviously with some disagreements, along two lines. The one approach that became widely adopted is the Wright-Fisher (WF) model (Crow and Kimura 1970; Hartl and Clark 1997), formalized by Fisher (1930) and Wright (1931). They define genetic drift to be solely due to the random sampling of genes. Thus, in generation transition, the variance in the frequency of a variant allele, x, would be

Revised version: Hagedoorn and Hagedoorn (1921) may be the first to suggest that random forces (accidents, for example) could impact long-term evolution. In the same decade, models of genetic drift were developed, obviously with some disagreements, along two lines. The one approach that became widely adopted is the Wright-Fisher (WF) model (Crow and Kimura 1970; Hartl and Clark 1997), formalized by Fisher (1930) and Wright (1931). They define genetic drift to be solely due to the random sampling of genes. Thus, in generation transition, the variance in the frequency of a variant allele, x, would be where N is the population size. The effective population size, Ne, is a function of varying N’s in space, time and demography. For example, when N fluctuates, Ne would be the harmonic mean of N’s across time (Crow and Kimura 1970; Caballero 1994; Charlesworth 2009). Weakly or strongly, Ne should still reflect the actual population size, N. In the WF model, genetic drift is governed by 1/ Ne in haploid populations, or 1/(2Ne) in diploid populations.

Removed in the revised version: where N is the population size. The effective population size, Ne, is a function of varying N’s in space, time and demography. For example, when N fluctuates, Ne would be the harmonic mean of N’s across time. Weakly or strongly, Ne should still reflect the actual population size, N. In the WF model, genetic drift is governed by 1/ Ne in haploid populations, or 1/(2Ne) in diploid populations.

Earlier version: At the same time, JBS Haldane introduced the branching process to model genetic drift (Haldane 1927; Haldane 1932). In the Haldane model, each gene copy is independently transmitted to K descendants with the mean and variance of E(K) and V(K). Subsequently, the model was expanded by Kimura and Crow (1963) and others (Wright 1969; Gillespie 1975; Eldon and Wakeley 2006; Chen, et al. 2017; Sackman, et al. 2019). Genetic drift is simply V(K), i.e., the variance in transmission success among neutral variants. There would be no drift if V(K) = 0. To scale drift at the population level, the variances of changes in gene frequency is governed by V(K)/N (or V(K)/Ne; Kimura and Crow (1963); Chen, et al. (2017); see also the Supplementary Information). In short, V(K) is the reproductive component of genetic drift while 1/N is the sampling component or the scaling factor.

Revised version: At the same time, JBS Haldane introduced the branching process to model genetic drift (Haldane 1927; Haldane 1932). In the Haldane model, each gene copy is independently transmitted to K descendants with the mean and variance of E(K) and V(K). In its essence, genetic drift is V(K), i.e., the variance in transmission among neutral variants. There would be no drift if V(K) = 0. To scale drift at the population level, the variance becomes V(K)/N (Kimura and Crow (1963); Chen, et al. (2017); see also the Supplementary Information). In short, V(K) is the reproductive component of genetic drift while 1/N is the sampling component or the scaling factor.

Added in the revised version: It should be noted that the WF model has been modified in many directions (Wright 1969; Gillespie 1975; Eldon and Wakeley 2006; Chen, et al. 2017; Sackman, et al. 2019) and, in particular, with the inclusion of V(K) (Kimura and Crow 1963; Sjodin, et al. 2005; Der, et al. 2011; Cannings 2016). These modifications share the basic features as well as the difficulties in resolving the paradoxes encountered by the WF model (see Results). Such difficulties stem from the sampling step that is tied to a pre-determined population size. These difficulties are shared by other models with population sampling, such as the Moran model (Moran 2008). In Discussion, we will present an overview of the modified WF models where these paradoxes will be discussed in relation to both the standard and modified WF models.

Earlier version: In the WF model, drift is governed only by 1/N (or 1/Ne). By dropping the reproductive component (see Eqs. 2-3 for further discussions), it has led to many paradoxes. A most curious one is genetic drift in relation to changing population size. The WF model dictates stronger drift when N is smaller. However, when N is very small (say, N = 10 in a bacteria or yeast population), that increases to 20, 40, 80 and so on, there is in fact little drift in this exponential phase. Drift will increase when N grows to near the carrying capacity. This trend is the exact opposite of the main property of the WF model.

Revised version: Among the paradoxes, a most curious one is genetic drift in relation to changing population size. The WF model dictates stronger drift when N is smaller. However, when N is very small (say, N = 10 in a bacteria or yeast population) and increases to 20, 40, 80 and so on, there is in fact little drift in this exponential phase. Drift will intensify when N grows to near the carrying capacity. This trend is the exact opposite of the tenet of the WF model. The second paradox concerns the genetic drift of sex-linked genes in relation to sex-dependent breeding successes (Charlesworth and Charlesworth 2000; Bachtrog 2013; Cortez, et al. 2014; Wilson Sayres, et al. 2014; Makova, et al. 2024). A third paradox is about drift strength when selection is also at work. A new mutation with a selective advantage would always be fixed if there is no random drift. With drift, fixation is probabilistic but the probability does not depend on N or Ne. This curious property echoes the view that N or Ne is a scaling factor of drift, which is determined mainly by V(K) (see Results).

Removed in the revised version: The second paradox concerns the genetic drift of sex chromosomes. When normalized by their copy numbers, genetic diversities still differ among autosomes, X and Y in mammals (Hammer, et al. 2001; Wang, et al. 2014). Although selection has been frequently invoked (Charlesworth and Charlesworth 2000; Bachtrog 2013; Cortez, et al. 2014; Wilson Sayres, et al. 2014; Makova, et al. 2023), the issue is if genetic drift due to V(K) differences between sexes is a sufficient explanation. A third paradox is about drift strength when selection is also at work. A new mutation with a selective advantage of s would always be fixed if there is no random drift. The fixation probability with genetic drift is 2s/V(K) (see Results). Drift under selection is hence dependent on V(K) but, somewhat non-intuitively, not on N.

Earlier version: The power of the Haldane model in resolving these paradoxes is the focus of this study. Furthermore, the companion study (Wang et al. 2024) will address a larger issue - the evolution of biological systems. An example is evolution in multi-copy gene systems, whereby evolution proceeds in two stages - between as well as within individuals. Multi-copy gene systems include viruses, mitochondria, transposons and multi-copy genes such as ribosomal RNA (rRNA) repeat units, which are conceptually the WF model. Strictly speaking, even diploidy is also a multi-copy system as a diploid individual usually produces large quantities of gametes that compete to be transmitted (Silver 1985; Wu, et al. 1988; Wu, et al. 1989; Lindholm, et al. 2016; Courret, et al. 2019). In discussion, we will further explore systems of biological evolution that are beyond the power of the WF Model. They include the evolution under domestication, somatic cell evolution and stochastic species competition.

Revised version: The power of the Haldane model in resolving these paradoxes is the focus of this study. Furthermore, the companion study (Wang et al. 2024) addresses a fourth paradox – the evolution of multi-copy gene systems, whereby evolution proceeds in two stages - between as well as within individuals. Multi-copy gene systems including viruses, mitochondria, transposons and ribosomal genes, are hence beyond the power of the WF models. Strictly speaking, even diploidy is also multi-copy systems (Silver 1985; Wu, et al. 1988; Wu, et al. 1989; Lindholm, et al. 2016; Courret, et al. 2019).

Added in the revised version: In this study, we present the integration of the Haldane and WF models into the WFH model. It is essentially the extension of the Haldane model with the capacity for N regulation. In this integration, the wealth of results obtained by the WF model are often adequate approximations for the more exact results. Most important of all, the WFH model can resolve the many paradoxes of molecular evolution.

Removed in the revised version: Through theory development supported by experimental data, we present the integration of the Haldane and the WF models into the WFH model. The integrated model is essentially the Haldane model with the approximate solutions obtained by the WF model. Cases where the WF model is not operative will still be handled by the Haldane model as in multi-copy gene systems. In short, the WF model can be considered special cases of the integrated WFH model. Therefore, the wealth of results obtained by the WF model are usually adequate approximations. Nevertheless, paradoxes do emerge as presented in this study.

Results

Earlier version: In the WFH model, genetic drift is defined as in the Haldane model of branching process (see Eq. (4) of Chen, et al. (2017). Let K be the number of progenies receiving the gene copy of interest from a parent; K would follow the same distribution for all neutral variants. In haploids, K is equivalent to the progeny number of each parent. In diploids, K is the transmission success of each gene copy, which is separately tracked. By using 2N for the gene copy number in the population, the WF as well as Haldane model treats each diploid individual as two-haploids. By the Taylor approximations of the ratio of two variables (Kendall, et al. 2006), we obtain the approximation as

Revised version: In the WFH model, genetic drift is defined as in the Haldane model of branching process (see Eq. (4) of Chen, et al. (2017)). Let K be the number of progenies receiving the gene copy of interest from a parent; K would follow the same distribution for all neutral variants. In haploids, K is equivalent to the progeny number of each parent. In diploids, K is the transmission success of each gene copy, which is separately tracked. By using 2N for the gene copy number in the population, the WF as well as Haldane model treats each diploid individual as two-haploids. By the Taylor approximations of the ratio of two variables (Kendall, et al. 2006), we obtain the approximation as where x is the variant frequency (see Supplementary Information). This is the same approximation, when V(K) is not very large, as obtained by the WF model (Kimura and Crow 1964; see Chen et al. 2017 for details). For Eq. (2), we provide further extensive simulations on the accuracy in both the fixation probability and fixation time of neutral variants (see Supplementary Information). The equivalence makes a simple point that the WF model represent special cases of the WFH model, and the results are often, but not always, shared.

Removed in the revised version: where x is the variant frequency (see Supplementary Information). This is the same approximation, when V(K) is not very large, as obtained by the WF model (Kimura and Crow 1964; see Chen et al. 2017 for details). For Eq. (2), we provide further extensive simulations on the accuracy in both the fixation probability and fixation time of neutral variants (see Supplementary Information). The equivalence makes a simple point that the WF model represent special cases of the WFH model, and the results are often, but not always, shared.

In the WFH model, there would be no genetic drift if V(K) = 0, regardless of the value of N. When E(K) = 1 and N stays constant, we obtain

Earlier version: While Eqs. 2 and 3 are biologically compatible with the Haldane model, their application to the WF model, as has been done(Kimura and Crow 1963; Crow and Denniston 1988; Chen, et al. 2017), is dubious. This is because the sampling scheme of the WF model dictates Poisson distribution for K, hence, V(K) = E(K). To permit V(K) ≠ E(K), the WF model would have to be modified to resemble the Haldane model(Kimura and Crow 1963; Caballero 1994). Such an awkward “fix” in fact defeats the reason for adopting the simpler WF model over the Haldane model

Revised version: Eq. (2) or Eq. (3) as applied to the WF model (Kimura and Crow 1963; Crow and Denniston 1988; Chen, et al. 2017) is biologically dubious. This is because the sampling scheme of the WF model dictates Poisson distribution for K, hence, V(K) = E(K). To permit V(K) ≠ E(K), the WF model would have to be modified to resemble the Haldane model (Kimura and Crow 1963; Caballero 1994). In Discussion, we shall analyze this “fix” of the WF model that still fails to resolve the paradoxes raised in this study.

Earlier version: In Table 1, the data of progeny number from the literature show over-dispersion in all taxa surveyed, i.e., V(K) > E(K). (Strictly speaking, the progeny number is the sum of two numbers from the same distribution of K; nevertheless, V(K) > E(K) is a valid interpretation of the data.) In fact, V(K) > 10 E(K) can often be found, e.g., the ratio of V(K)/E(K) among males of the mandrill (Mandrillus sphinx), the great reed warbler (Acrocephalus arundinaceus), and the rhesus macaque (Macaca mulatta) at 19.5, 12.6 and 11.3, respectively (Hasselquist 1995; Setchell, et al. 2005; Dubuc, et al. 2014). We have also found that V(K) may be far larger than E(K) in other biological systems such as viruses (Ruan, Luo, et al. 2021; Ruan, Wen, et al. 2021; Hou, et al. 2023).

Revised version: In Table 1, the data of progeny number from the literature indeed show over-dispersion in all taxa surveyed, i.e., V(K) > E(K). In fact, V(K) > 10 E(K) can often be found, e.g., the ratio of V(K)/E(K) among males of the mandrill (Mandrillus sphinx), the great reed warbler (Acrocephalus arundinaceus), and the rhesus macaque (Macaca mulatta) at 19.5, 12.6 and 11.3, respectively (Hasselquist 1995; Setchell, et al. 2005; Dubuc, et al. 2014). We have also found that V(K) may be far larger than E(K) in other biological systems such as viruses (Ruan, Luo, et al. 2021; Ruan, Wen, et al. 2021; Hou, et al. 2023).

The paradox of genetic drift in relation to N changes (Paradox of changing N)

Empirical demonstration and simulation

This “paradox of changing N” is that “genetic drift increases when N increases”. Fig. 1a shows the results in a cell culture in the exponential growth phase where each cell doubles every 13 hours. Hence, V(K) ∼ 0 with almost no drift when N is as small as < 50. Genetic drift is expected to increase as N approaches the carrying capacity. The paradox is exemplified through computer simulations in Fig. 1b-c. These panels show the drift pattern in the Haldane model and the WF model, respectively, when the populations are growing as the logistic growth model (see Methods). In the WF model, the drift is strong initially but weakened as N increases. The pattern is reversed in the Haldane model. The contrast between the two panels is the paradox.

The paradox of genetic drift when the population size (N) changes.

(a) In the laboratory cell culture, nearly all cells proliferate when N is very small (interpreted in the next panel). (b) The simulation of Haldane model shows little genetic drift in the exponential phase. As in (a), drift may increase due to the heightened competition as the population grows. (c) Simulation by the WF model shows a pattern of drift opposite of the Haldane model. (d) The patterns of drift at the low and high N are analyzed in the framework of the logistic growth model. (e-f) Measurements of genetic drift in laboratory yeast populations at low and high density as defined in (d). The progeny number of each cell, K, is counted over 4 or 5 intervals as shown by the dots, the sizes of which reflect the cell number. E(K) and V(K) are presented as well. In panel (f), the change of offspring number overtime, denoted as V(ΔK), is shown above the braces. (g) The variance of offspring number V(K) increases, observed in (e) and (f), as population size increases.

To verify the simulations of Fig. 1b-c, we study the growth of yeast populations under laboratory conditions where cells appear to increase following the logistic growth curve. As shown in Fig. 1d, the lower portion of the S-curve portrays the exponential growth in a low-density population, while the upper curve indicates a slowdown in growth near the population’s carrying capacity. We directly recorded the individual output (K) of each yeast cell in the early (n = 25) and late (n = 65) stage of population growth under real-time high-content fluorescence imaging (Methods).

It’s evident that V(K) is decoupled from E(K). In the early stage, E(K) is 10.04 with V(K) = 0.60 over five one-hour intervals. In the high-density phase, E(K) decreases to 2.83 but V(K) increases to 1.09 in four intervals (Fig. 1e-f). Most interestingly, during the late stage of population growth, V(ΔK) (adjusted to the four-hour time span for comparison; see Methods) indeed increases as N increases (Supplementary Table 3). At the last time point, the population is closest to the carrying capacity and V(ΔK) experiences a substantial leap from around 1.12 to 1.82. Overall, the variance in progeny number V(K) in a high-density population is always greater that in a low-density population (Fig. 1g). Indeed, V(K) may exhibit an increase with the increase in N, sometimes outpaces the latter.

Measurements of these two growth phases demonstrate the “paradox of changing N” by showing i) there is almost no drift when N is very small; ii) as N increases, V(K), hence, genetic drift would also increase. Furthermore, V(K) may increase sharply when N is close to the carrying capacity. This trend appears to echo the field observations by Coltman, et al. (1999) who show that V(K) increases more than 5-fold when N increases only 2.5-fold in the United Kingdom (UK) population of Soay rams (Ovis aries).

The density-dependent Haldane (DDH) model

The “paradox of changing N” can be explained by Ne = N/V(K) if the numerator and denominator move in the same direction. Most interestingly, V(K) may increase more than N itself. This has been shown when V(K) is close to zero while N is growing exponentially. In addition, when N is very near the carrying capacity, even a small increase in N would intensify the competition and inflate V(K). This is shown in the yeast experiment of Fig. 1d and supported by Coltman, et al. (1999) field study of ram populations in UK. Thus, the relationship between N and V(K) may realistically lead to this paradox.

Removed in the revised version: We now extend the Haldane model by incorporating density-dependent regulation of N in order to define the conditions of this paradox. The model developed here is on an ecological, rather than an evolutionary time scale. The simplest model of N regulation is the logistic equation:

Earlier version: with Ck being the carrying capacity. In the ecological time scale, a changing N would mean that the population is departing from Ck or moving toward it (Fig. 2a). (As will be addressed in Discussion, changing N in the WF model is depicted in Fig. 2b whereby N is at Ck. Ck may evolve too, albeit much more slowly in an evolutionary time scale.) Here, we consider changes in N in the ecological time scale. Examples include the exponential growth when N ≪ Ck, seasonal fluctuation in N (Frankham 1995; Charlesworth 2009) and competition modeled by the Lotka-Volterra equations (Bomze 1983).

Revised version: We now extend the Haldane model by incorporating density-dependent regulation of N in order to define the conditions of this paradox. The model developed here is on an ecological, rather than an evolutionary time scale. The simplest model of N regulation is the logistic equation: with Ck being the carrying capacity. In the ecological time scale, a changing N would mean that the population is departing from Ck or moving toward it (Fig. 2a). (As will be addressed in Discussion, changing N in the WF model is depicted in Fig. 2b whereby N is at Ck. Ck may evolve too, albeit much more slowly in an evolutionary time scale.) Here, we consider changes in N in the ecological time scale. Examples include the exponential growth when N ≪ Ck, seasonal fluctuation in N (Frankham 1995; Charlesworth 2009) and competition modeled by the Lotka-Volterra equations (Bomze 1983).

The meaning of population size (N) changes in ecology vs. in population genetics.

(a) In ecology, changing N would generally mean a population approaching or departing the carrying capacity. (b) In population genetics, a population of size N is assumed to be at the carrying capacity, Ck. Thus, changes in N would mean an evolving Ck, likely the consequence of environmental changes. The arrows indicate the disparity in time scale between the two scenarios.

Earlier version: Details of the DDH model are given in the Supplementary Information. A synopsis is given here: We consider a haploid population with two neutral alleles. The population size at time t is Nt. It is assumed that E(K) > 1 when N < Ck and E(K) < 1 when N > Ck as defined in Eq. (5) below.

Revised version: Details of the DDH model are given in the Supplementary Information. A synopsis is given here: We consider a non-overlapping haploid population with two neutral alleles. The population size at time t is Nt. We assume that expected growth rate E(K) is greater than 1 when N < Ck and less than 1 when N > Ck, as defined by Eq. (5) below:

Earlier version: The slope of E(K) vs. N, as shown in Fig 3a, is a function of z. To obtain V(K), we assume that K follows the negative binomial distribution whereby parents would suffer reproduction-arresting injury with a probability of pt at each birthing (Supplementary Information). V(K) can then be expressed as

Revised version: The slope of E(K) vs. N (i.e., the sensitive of growth rate to changes in population size), as shown in Fig 3a, depends on z. To determine the variance V(K), we assume that K follows the negative binomial distribution whereby parents would suffer reproduction-arresting injury with a probability of pt at each birthing (Supplementary Information). Accordingly, V(K) can then be expressed as

Added in the revised version: By Eq. (6), the ratio of V(K)/E(K) could be constant, decrease or increase with the increase of population size. With E(K) and V(K) defined, we could obtain the effective population size by substituting Eq. (5) and Eq. (6) into Eq. (3).

Added in the revised version: Eq. (7) presents the relationship between effective population size (Ne) and the population size (N) as shown in Fig. 3. The density-dependent E(K) could regulate N with different strength (Fig. 3a). The steeper the slope in Fig. 3a, the stronger the regulation.

Genetic drift as a function of population size in the DDH model.

For all panels, the carrying capacity is Ck = 10,000 and the intrinsic growth rate is r = 2. (a) When N increases, E(K) decreases as modeled in Eq. (5). The z value of Eq. (5) (0.1, 1.5 and 3) determines the strength of N regulation, indicated by the slope of E(K) near Ck = 10,000. (b) Depending on the strength of N regulation near Ck, genetic drift can indeed decrease, increase or stay nearly constant as the population size increases. Thus, the conventional view of Ne being positively dependent on N is true only when the regulation of N is weak (the green line). At an intermediate strength (the red line), Ne is nearly independent of N. When the regulation becomes even stronger at z = 3, Ne becomes negatively dependent on N. (c-e) V(K)/E(K) of Eq. (6)) is shown as a function of N. The results of panel (b) are based on a constant V(K)/E(K) shown in panel (c). Interestingly, the results of panel (b) would not be perceptibly changed when V(K)/E(K) varies, as shown in panel (d) and (e).

Removed in the revised version: By Eq. (6), the ratio of V(K)/E(K) could be constant, decrease or increase with the increase of population size. With E(K) and V(K) defined, we could obtain the effective population size by substituting Eq. (5) and Eq. (6) into Eq. (3).

Removed in the revised version: Eq. (7) presents the relationship between effective population size (Ne) and the population size (N) as shown in Fig. 3. The density-dependent E(K) could regulate N with different strength (Fig. 3a). The steeper the slope in Fig. 3a, the stronger the regulation.

The main results of this study are depicted in Fig. 3b. First, with no or little regulation of N, Ne and N are strongly correlated. The green dashed lines portray the conventional view of decreasing drift with increasing N. Second, under a particular strength of N regulation, the red line shows no dependence of Ne on N, meaning that genetic drift is independent of N. Finally, as N becomes strongly regulated, Ne and N would be negatively correlated as the blue dashed line shows. This trend is the paradox of changing N.

In summary, genetic drift effect can indeed decrease, increase or stay nearly constant as the population size increases, depending on the strength of N regulation near Ck. We further note that the V(K)/E(K) ratio of Eq. (7) can be independent (Fig. 3c), positively dependent (Fig. 3d), or negatively dependent (Fig. 3e) on N. Interestingly, the results of Fig. 3b are nearly identical when the V(K)/E(K) ratio changes (Supplementary Figs 1-3). These results show that the strength of genetic drift depends on the ecology that governs E(K), V(K) and N.

In the WFH framework, the paradox of changing N may not be an exception but a common rule in natural populations. Note that the WF approximation yielding Eq. (2) assumes nearly constant N. The assumption would imply very strong N regulation near Ck and, according to the DDH model, this is the condition for realizing the paradox of changing N. Coltman et al. (1999)’s observations of the reproduction in rams may be such a realization.

The paradox of genetic drift in sex chromosomes

Earlier version: Since the relative numbers of Y, X and each autosome (A) in the human population are 1:3:4, the WF model would predict the genetic diversity (θY, θX and θA, respectively) to be proportional to 1:3:4. In a survey of human and primate genetic diversity, θY is almost always less than expected as has been commonly reported (Hammer, et al. 2001; Wang, et al. 2014; Wilson Sayres, et al. 2014; Makova, et al. 2023). The low θY value has been used to suggest that human Y chromosomes are under either positive or purifying selection (Wang, et al. 2014; Wilson Sayres, et al. 2014; Makova, et al. 2023). Similarly, the reduced diversity X-linked genes was interpreted as a signature of positive selection (Pan, et al. 2022).

Revised version: Since the relative numbers of Y, X and each autosome (A) in the human population are 1:3:4, the WF model would predict the genetic diversity (θY, θX and θA, respectively) to be proportional to 1:3:4. In a survey of human and primate genetic diversity, θY is almost always less than expected as has been commonly reported (Hammer, et al. 2001; Wang, et al. 2014; Wilson Sayres, et al. 2014; Makova, et al. 2024). The low θY value has been used to suggest that human Y chromosomes are under either positive or purifying selection (Wang, et al. 2014; Wilson Sayres, et al. 2014; Makova, et al. 2024). Similarly, the reduced diversity X-linked genes was interpreted as a signature of positive selection (Pan, Liu, et al. 2022).

As pointed out above, under-estimation of genetic drift is a major cause of over-estimation of selective strength. Our goal is hence to see if genetic drift of different strength between sexes can account for the observed genetic diversities. Details will be presented in Wang et al. (2024). Below is a synopsis.

Let Vm and Vf be the V(K) for males and females respectively and α’ = Vm/Vf is our focus. (We use α’ for the male-to-female ratio in V(K) since α is commonly used for the ratio of mutation rate (Miyata, et al. 1987; Makova and Li 2002).) The three ratios, θY /θX (denoted as RYX), θY /θA(RYA) and θX /θA (RXA), could be expressed as the functions of α’, which incorporate the relative mutation rates of autosomes, X and Y (Supplementary Information). These relative mutation rates can be obtained by interspecific comparisons as done by Makova and Li (2002).

We note that there are many measures for the within-species genetic diversity, θ = 4Neμ. Under strict neutrality and demographic equilibrium, these measures should all converge. In the neutral equilibrium, the infinite site model dictates the frequency spectrum to be ξi = θ/i, where ξi is the number of sites with the variant occurring i times in n samples. Since every frequency bin is a measure of θ, different measures put different weights on the i-th bin (Fay and Wu 2000; Fu 2022). While Π, the mean pairwise differences between sequences, is most commonly used in the literature, we use several statistics to minimize the possible influences of selection and demography (Wang et al. 2024). In this synopsis, we used the Watterson estimator (Watterson 1975) as the measure for θY, θX and θA by counting the number of segregating sites.

Earlier version: Using any of the three ratios (θY /θX, θY /θA and θX /θA), we can obtain α’; for example, RYA = θY /θA and, hence,

Revised version: Using any of the three ratios (θY /θX, θY /θA and θX /θA), we can obtain α’; for example, RYA = θY /θA and, hence, where y is the mutation rate of Y-linked sequences relative to autosomal sequences. With rearrangement,

Removed in the revised version: where y is the mutation rate of Y-linked sequences relative to autosomal sequences. With rearrangement,

Similar formulae of α’ can be obtained for RYX and RXA but the accuracy for estimating α’ is highest by RYA whereas RXA is the least accurate for this purpose (Supplementary Information).

Table 2 presents the α’ estimates in chimpanzees and bonobos. This is part of a general survey in mammals with a strong emphasis on primates and humans (Wang et al. 2024). It is almost always true that α’ > 5 in primates. Sources of data used are also given in Table 2. As shown for chimpanzees, α’ is often far larger than 5, above which the resolution is very low as can be seen in Eq. (8). Note that, when RYA is under-estimated and approaching y/8, α’ would increase rapidly when the denominator is close to 0. That is why α’ often becomes infinity (Supplementary Fig. 5). The estimated α’ (MSE) in Table 2 alleviates the problem somewhat by using all three ratios to calculate the mean square error (MSE) (Supplementary Information).

Among primates surveyed, bonobo is the only exception with α’ < 1. While chimpanzees and bonobos are each’s closest relatives, their sexual behaviors are very divergent (de Waal 1995; De Waal and Lanting 2023). With unusually strong matriarchs, bonobo society seems to stand out among primates in sexual dominance. The α’ value is important in behavioral studies and, particularly, in primatology and hominoid research. If 0.1% of males sire 100 children and the rest have the same K distribution as females, Vm/Vf would be ∼ 10. Such outlier contribution is the equivalent of super-spreaders in viral evolution which may easily be missed in field studies. In this brief exposition, we highlight again that V(K) or, more specifically, Vm and Vf are key to genetic drift.

The paradox of genetic drift under selection

Genetic drift operates on neutral as well as non-neutral mutations. Let us assume a new mutation, M, with a frequency of 1/N is fixed in the population with the probability of Pf. Fisher (1930) first suggested that the fixation probability of an advantageous mutation should increase in growing populations, while decrease in shrinking populations. If there is no genetic drift (i.e., infinite population size in WF model), a beneficial mutation will always be fixed and Pf = 1. In the WF model, it is well known that Pf ∼ 2s for a new advantageous mutation, with fitness gain of s, while population size is large (Haldane 1932; Kimura 1962). This seems paradoxical that the determinant of genetic drift, 1/Ne, does not influence Pf (Otto and Whitlock 1997; Lanfear, et al. 2014).

Added in the revised version: In the Method section, we show that Pf under the Haldane model is given by where M0 and W0 are the initial number of mutant allele (M allele) and wildtype allele (W allele), with N = M0 + W0. We note that uM(t) and uW(t), respectively, are the extinction probabilities of M allele and W allele by generation t with the initial number of alleles of 1. While obtaining the direct analytical solution of Eq. (9) may not be feasible, we could obtain its numerical solution due to its convergence as t increases (see Methods). The accuracy of the numerical solution from Eq. (9) is confirmed through simulation (Supplementary Fig. 4). Moreover, with the aid of the WF model (see Supplementary Information), we could obtain the approximation of Eq. (9) as follows.

Added in the revised version: We verify Eq. (10) by both the numerical solution from Eq. (9) and simulations based on the branching process (Supplementary Fig. 4). The fixation probabilities obtained by numerical solution vs. those inferred from Eq. (10) are shown in Fig. 4. The salient feature is that the fixation probability of 2s, as in the classical formula, would be a substantial over-estimate when V(K) is larger than E(K). Eq. (10) is sufficiently accurate as long as N≥50. When N is as small as 10, the theoretical result is biased. Indeed, at such a low N value, the population is prone to extinction. The DDH model presented above should rectify this deficiency.

Fixation probability of a new advantageous mutation in the Haldane model.

The fixation probabilities of a new advantageous mutation with the selective advantage of s = 0.1 are calculated based on approximate solution from Eq. (9) (i.e., 2s/V(K)) as well as numerical solution from Eq. (10). The numerical solution from Eq. (10) has been confirmed accurate by simulations (Supplementary Fig. 4). (a-b) When N < 50, the approximate fixation probability (the gray line) is lower than the simulated values (the color lines) due to population extinction. (c-d) By the Haldane model, the expected fixation probability of Eq. (9) is accurate when N reaches 100, as in most natural populations.

Added in the revised version: The main message of Eq. (10) is that genetic drift under positive selection is influenced by s and V(K) but not by N. This independence from N seems intuitively obvious: when an advantageous mutation increases to say, 100 copies in a population (depending mainly on s), its fixation would be almost certain, regardless of N. This may be a most direct argument against equating genetic drift with N, or Ne (which is supposed to be a function of N’s).

Removed in the revised version: In the Method section, we show that Pf under the Haldane model is given by

Removed in the revised version: where M0 and W0 are the initial number of mutant allele (M allele) and wildtype allele (W allele), with N = M0 + W0. We note that uM(t) and uW(t), respectively, are the extinction probabilities of M allele and W allele by generation t with the initial number of alleles of 1. While obtaining the direct analytical solution of Eq. (9) may not be feasible, we could obtain its numerical solution due to its convergence as t increases (see Methods). The accuracy of the numerical solution from Eq. (9) is confirmed through simulation (Supplementary Fig. 4). Moreover, with the aid of the WF model (see Supplementary Information), we could obtain the approximation of Eq. (9) as follows.

Removed in the revised version: We verify Eq. (10) by both the numerical solution from Eq. (9) and simulations based on the branching process (Supplementary Fig. 4). The fixation probabilities obtained by numerical solution vs. those inferred from Eq. (10) are shown in Fig. 4. The salient feature is that the fixation probability of 2s, as in the classical formula, would be a substantial over-estimate when V(K) is larger than E(K). Eq. (10) is sufficiently accurate as long as N≥50. When N is as small as 10, the theoretical result is biased. Indeed, at such a low N value, the population is prone to extinction. The DDH model presented above should rectify this deficiency. The main message of Eq. (10) is that genetic drift under positive selection is influenced by V(K) but not by N. Hence, V(K) is a vital factor in determining genetic drift, certainly no less important than 1/N.

Discussion

Earlier version: Genetic drift should broadly include all random forces affecting evolutionary trajectories. Many of these forces are excluded from the WF model, thus resulting in the many paradoxes addressed in the two companion papers, as well as the over-estimation of selection in the literature.

Revised version: Genetic drift should broadly include all random forces affecting evolutionary trajectories. Many of these forces are excluded from the WF model, thus resulting in the many paradoxes addressed in the two companion papers, as well as the over-estimation of selection in the literature. We shall now dissect these paradoxes in detail in relation to the standard and modified WF models.

Earlier version: Compared with the conventional WF model of genetic drift, the Haldane model is conceptually more intuitive as well as more inclusive in depicting the real populations. However, analytical solutions of the branching process can be more challenging to obtain. The WF model is often more tractable mathematically. For example, Pf ∼ 2s/V(K) is a good WF approximation (Kimura and Crow 1963) to the Haldane model when N is no smaller than 100 (Fig. 4). The integrated WFH model is therefore conceptually appealing and computationally feasible.

Revised version: Compared with the WF models of genetic drift, the Haldane model is more intuitive as well as inclusive in depicting real populations. However, analytical solutions of the branching process can be more challenging to obtain. The WF model is often more tractable mathematically. For example, Pf ∼ 2s/V(K) is a good WF approximation (Kimura and Crow 1963) to the Haldane model when N is above 100 (Fig. 4). The integrated WFH model is hence both biologically realistic and analytically tractable (or at least computationally feasible).

Earlier version: Nevertheless, in many circumstances where the WF model is not operative, genetic drift defined by the Haldane model would demand new solutions. The DDH model is such an example, developed to resolve the paradox of changing N’s. The DDH model presents the conditions under which the WF model arrives in the approximate solution (i.e., Eq. (2)). The substitution of E(N) for N makes the approximation possible but effectively requires strong regulation of N. Strong regulation is the main condition for the paradox of changing N’s. To arrive at the approximation, the WF model in fact makes an assumption that contradicts its own prediction (that drift becomes weaker as N increases).

Revised version: Nevertheless, in many circumstances where the WF model is not operative, genetic drift defined by the Haldane model would demand new solutions. The DDH model is such an example, developed to resolve the paradox of changing N’s. The DDH model presents the conditions under which the WF model arrives in the approximate solution (i.e., Eq. (2)). The substitution of E(N) for N makes the approximation possible but effectively requires strong regulation of N. Interestingly, if N is strongly regulated, genetic drift is positively correlated with N, opposite of the WF model prediction (see Fig. 3). In short, the Haldane model can incorporate population size regulation, which is beyond the power of the WF model.

Earlier version: We shall now clarify the “paradox of changing N’s” on two fronts. First, the DDH model is on an ecological time scale whereby a population is either approaching or departing the carrying capacity (Fig. 2a). In many studies, the population size changes in the evolutionary time scale (Fig. 2b). In such studies, it is the carrying capacity Ck that is evolving while N hovers around Ck. For example, by the PSMC model (Li and Durbin 2011), both the European and Chinese populations experience a severe bottleneck 10-60 kyr ago. During that Last Glacial Maximum, environmental harshness may have reduced Ck drastically. Averaged over the long-time span, N should be at or very near Ck. Since Ck presumably evolves slowly, N can be assumed constant over long stretches of time. Later models with different powers than PSMC have been proposed to study N’s (or Ck’s) in the evolutionary time scale (Hu, et al. 2023). Second, Gillespie has proposed the concept of “genetic draft” that may also reverse the effect of population size on the rate of substitution due to selection on linked variations (Gillespie 2000, 2001). Genetic draft of Gillespie is thus distinct from the strictly neutral process discussed in this study.

Revised version: We shall further clarify the “paradox of changing N’s” on two fronts. First, the DDH model is on an ecological time scale whereby a population is either approaching or departing the carrying capacity (Fig. 2a). However, in many studies, the population size changes in the evolutionary time scale (Fig. 2b). For example, by the PSMC model (Li and Durbin 2011), both the European and Chinese populations experience a severe bottleneck 10-60 kyr ago. Presumably, various environmental forces may have reduced Ck drastically. Averaged over the long-time span, N should be at or very near Ck. Second, Gillespie has proposed the concept of “genetic draft” that may also reverse the effect of population size on the rate of substitution due to selection on linked variations (Gillespie 2000, 2001). Genetic draft of Gillespie is thus distinct from the strictly neutral process discussed in this study.

Added in the revised version: The standard WF model has been extended in several directions (overlapping generations, multiple alleles, ploidy, etc.). The modification most relevant to our studies here is the introduction of V(K) into the model, thus permitting V(K) ≠ E(K). While the modifications are mathematically valid, they are often biologically untenable. Kimura and Crow (1963) may be the first to offer a biological mechanism for V(K) ≠ E(K), effectively imposing the Haldane model on the WF model. Other models (Kimura and Crow 1963; Lynch, et al. 1995; Sjodin, et al. 2005; Der, et al. 2011; Cannings 2016) indeed model mathematically the imposition of the branching process on the population, followed by the WF sampling. The constructions of such models are biologically dubious but, more importantly, still unable to resolve the paradoxes. It would seem more logical to use the Haldane model in the first place by having two parameters, E(K) and V(K).

Added in the revised version: Even if we permit V(K) ≠ E(K) under the WF sampling, the models would face other difficulties. For example, a field biologist needs to delineate a Mendelian population and determine its size, N or Ne. In all WF models, one cannot know what the actual population being studied is. Is it the fly population in an orchard being sampled, in the geographical region, or in the entire species range? It is unsatisfactory when a population biologist cannot identify the population being studied. The Haldane model is an individual-output model (Chen, et al. 2017), which does not require the delineation of a Mendelian population.

Added in the revised version: We shall now review the paradoxes specifically in relation to the modified WF models, starting with the multi-copy gene systems such as viruses and rRNA genes covered in the companion study (Wang, et al. 2024). These systems evolve both within and between hosts. Given the small number of virions transmitted between hosts, drift is strong in both stages as shown by the Haldane model (Ruan, Luo, et al. 2021; Ruan, Wen, et al. 2021; Hou, et al. 2023). Therefore, it does not seem possible to have a single effective population size in the WF models to account for the genetic drift in two stages. The inability to deal with multi-copy gene systems may explain the difficulties in accounting for the SARS-CoV-2 evolution (Deng, et al. 2022; Pan, Liu, et al. 2022; Ruan, Wen, et al. 2022; Hou, et al. 2023; Ruan, et al. 2023).

Added in the revised version: We now discuss the first paradox of this study, which is about the regulation of N. In the general WF models, N is imposed from outside of the model, rather than self-generating within the model. When N is increasing exponentially as in bacterial or yeast cultures, there is almost no drift when N is very low and drift becomes intense as N grows to near the carrying capacity. As far as we know, no modifications of the WF model can account for this phenomenon that is opposite of its central tenet. In the general WF models, N is really the carrying capacity, not population size.

Added in the revised version: The second paradox of sex chromosomes is rooted in V(K) ≠ E(K). As E(K) is the same between sexes but V(K) is different, clearly V(K) = E(K) would not be feasible. The mathematical solution of defining separate Ne’s for males and females (Kimura and Crow 1963; Lynch, et al. 1995; Sjodin, et al. 2005; Der, et al. 2011; Cannings 2016) unfortunately obscures the interesting biology. As shown in Wang et al. (2024; MBE), the kurtosis of the distribution of K indicates the presence of super-breeder males. While the Haldane model can incorporate the kurtosis, the modified WF models are able to absorb only up to the variance term, i.e., the second moment of the distribution. The third paradox of genetic drift is manifested in the fixation probability of an advantageous mutation, 2s/V(K). As explained above, the fixation probability is determined by the probability of reaching a low threshold that is independent of N itself. Hence, the key parameter of drift in the WF model, N (or Ne), is missing. This paradox supports the assertion that genetic drift is fundamentally about V(K) with N being a scaling factor.

Removed in the revised version: A larger issue is that the WF model only covers a portion of biological evolution. All multi-copy gene systems where the evolution happens both within and between individuals, are out of reach. Viruses, transposons, mitochondria (Bazin, et al. 2006) and rRNA genes are all such systems. This topic will be addressed in the companion paper that focuses on the evolution of multi-copy rRNA genes. In fact, all diploids are the simplest kinds of the multi-copy gene systems. Because large quantities of gametes compete within individual to be transmitted, gametic selection, meiotic drive and genetic drift may all happen within individuals (Wu, et al. 1988; Wu, et al. 1989; Wu and Hammer 1991; Lindholm, et al. 2016; Courret, et al. 2019). The WF model is in essence a haploid model with 2N copies of genes rather than a true diploid model.

Earlier version: Like multi-copy gene systems, the WF model is limited, conceptually as well as operationally, by the difficulty in delineating WF populations. That is why multi-copy genes, as well as many other systems, do not fit the WF model, thus resulting in the lack of a genetic drift component in such studies. For example, domestications of animals and plants are very rapid processes of evolution (Diamond 2002; Larson and Fuller 2014; Purugganan 2019). Due to the very large V(K) in domestication, drift must have played a large role in such processes, but is rarely considered. Somatic cell evolution is another example with “undefinable” genetic drift (Wu, et al. 2016; Chen, et al. 2017; Chen, et al. 2019; Ruan, et al. 2020; Chen, Wu, et al. 2022). The Haldane (or WFH) model, as an “individual output” model, can handle these general cases of genetic drift. Finally, the WFH model is so general that even genetic drift and random forces in species competition can be incorporated.

Revised version: As the domain of evolutionary biology expands, many new systems do not fit into the WF models, resulting in the lack of a genetic drift component in their evolutionary trajectories. Multi-copy gene systems are obvious examples. Others include domestications of animals and plants that are processes of rapid evolution (Diamond 2002; Larson and Fuller 2014; Purugganan 2019; Chen, Yang, et al. 2022; Pan, Zhang, et al. 2022; Wang, et al. 2022). Due to the very large V(K) in domestication, drift must have played a large role. Somatic cell evolution is another example with “undefinable” genetic drift (Wu, et al. 2016; Chen, et al. 2017; Chen, et al. 2019; Ruan, et al. 2020; Chen, Wu, et al. 2022). The Haldane (or WFH) model, as an “individual output” model, can handle these general cases of genetic drift.

Added in the revised version: The Haldane model and the WF model are fundamentally different approaches to random forces of evolution. While the WF models encounter many biological contradictions, they have provided approximate mathematical solutions to more realistic scenarios. In systems such as in viral evolution (Ruan, Hou, et al. 2022; Hou, et al. 2023) or somatic cell evolution (Chen, Wu, et al. 2022; Zhai, et al. 2022) whereby the WF solution is absent, further development of the WFH model will be necessary.

Removed in the revised version: In conclusion, the integrated WFH model should clarify the biological meaning of genetic drift, wherever the WF and Haldane model can both be used. In those systems whereby only the Haldane model works, further development of the model will be necessary.

Methods

Cell culture and image analysis

Earlier version: NIH3T3 cells, a fibroblast cell line that was isolated from a mouse NIH/Swiss embryo, were stably transfected with the fluorescent, ubiquitination-based cell cycle indicator (Fucci) (Sakaue-Sawano, et al. 2008) plasmid using Lipofectamine 3000 Transfection Reagent (Invitrogen) following the manufacturer’s specified instructions. The Fucci-labeled cells exhibited distinct fluorescent signals indicative of G1, S, and G2/M phases, represented by red, yellow, and green, respectively. NIH3T3-Fucci cell was derived from single-cell colony and cultured in DMEM supplemented with 10% Calf Bovine Serum and penicillin/streptomycin. All cells were maintained at 37°C with 5% CO2. Subsequently, the cells underwent extended time-lapse imaging using high-content fluorescence microscopy (PerkinElmer Operetta CLS) equipped with a 10x objective lens. Images were captured hourly over a 100-hour period, and the analysis was conducted using ImageJ (Fiji) (Schneider, et al. 2012) to count the number of cells (Supplementary Table 1).

Revised version: NIH3T3 cells, a fibroblast cell line that was isolated from a mouse NIH/Swiss embryo, were stably transfected with the fluorescent, ubiquitination-based cell cycle indicator (Fucci) (Sakaue-Sawano, et al. 2008) plasmid using Lipofectamine 3000 Transfection Reagent (Invitrogen) following the manufacturer’s specified instructions. The Fucci-labeled cells exhibited distinct fluorescent signals indicative of G1, S, and G2/M phases, represented by red, yellow, and green, respectively. NIH3T3-Fucci cell was derived from single-cell colony and cultured in DMEM supplemented with 10% Calf Bovine Serum and penicillin/streptomycin. All cells were maintained at 37◻ with 5% CO2. Subsequently, the cells underwent extended time-lapse imaging using high-content fluorescence microscopy (PerkinElmer Operetta CLS) equipped with a 10x objective lens. Images were captured hourly over a 100-hour period, and the analysis was conducted using ImageJ (Fiji) (Schneider, et al. 2012) to count the number of cells (Supplementary Table 1).

Yeast strain construction

Strains were constructed on the genetic background of Saccharomyces cerevisiae strain BY4741. A GFP or BFP fluorescent protein expression cassette, under the control of the TDH3 promoter, was inserted into the pseudogene locus YLL017W. Transformations followed a published protocol (Gietz and Schiestl 2007). Transformants were plated on synthetic complete medium without uracil (SC-Ura), and from these, single colonies were selected. Confirmation of replacements was achieved through PCR, and cassette verification was performed using fluorescence microscopy. Subsequently, the constructed strains were cultivated non-selectively in YPD medium (1% Yeast extract, 2% Peptone, 2% D-glucose) at 30 ◻ on a rotary shaker.

Estimation of V(K) and E(K) in yeast cells

To discern division events, even under high concentration, we conducted co-cultures of the GFP-yeast and BFP-yeast at ratios of 1:1 and 1:25, with the initial cell concentration of 0.1% and 12.5%, respectively. Then the yeast cells were then continuously imaged under high-content fluorescence microscopy (PerkinElmer Operetta CLS) for 10 hours with 1-hour intervals to observe the individual offspring of GFP-yeast. Yeast cells with a distinct offspring number (K) within the initial 5 hours (with the high-density group limited to the first 4 hours) were documented (Supplementary Table 2). Subsequently, the mean and variance of K for each cell were calculated across various time intervals as follows.

Removed in the revised version: According to the law of total variance,

Earlier version: where It is the offspring number at time t for an initial single cell (i.e., the total number of progeny cells for a single cell after t hours), as documented in Supplementary Table 2. And Et-1, t represents the average of offspring number after single time interval (1 hour here) for each single cell from time t-1 to time t, while Vt-1, t is the corresponding variance. Utilizing the documented total cell count at time t (denoted by Nt), we could calculate the average of offspring number for a single yeast cell in a specific time interval, e.g., Et-1, t = Nt / Nt-1 and E[It-1] = Nt /N0. With some rearrangement,

Revised version: According to the law of total variance, where It is the offspring number at time t for an initial single cell (i.e., the total number of progeny cells for a single cell after t hours), as documented in Supplementary Table 2. And Et-1, t represents the average of offspring number after single time interval (1 hour here) for each single cell from time t-1 to time t, while Vt-1, t is the corresponding variance. Utilizing the documented total cell count at time t (denoted by Nt), we could calculate the average of offspring number for a single yeast cell in a specific time interval, e.g., Et-1, t = Nt / Nt-1 and E[It-1] = Nt /N0. With some rearrangement,

Applying the aforementioned equations, we observed that both Vt-1, t (i.e., V(K) within a one-hour interval) and V(It) (i.e., the V(K) from 0 to time t) for yeast cells under high density (12.5%) are generally greater than the values under low density (Supplementary Table 3). This suggests that the variance of offspring number in high-density conditions could be larger than that in a low-density context. It’s noteworthy that the estimated value of Vt-1, t could be less than zero (Supplementary Table 3), implying a very small variance in the number of offspring. Furthermore, the occurrence of negative values for Vt-1, t is more frequent under low density than high density, indicating the higher variance of offspring number in high density as previously suggested.

To show the variance of offspring number overtime, we also calculated the change of offspring number within successive one-hour intervals, denoted as V(It - It-1). This value is intricately linked to the variance of offspring overtime. To facilitate the comparison with the total variance from 0h to 4h, we set V(ΔK) = 4V(It - It-1) to account for the four-hour time span.

Simulation of genetic drift in the Haldane model and the Wright-Fisher (WF) model

Earlier version: In both models, the average number of offspring, denoted as E(Kt), follow Eq. (5), leading to the logistic population growth. The initial population size is set to 4, with initial gene frequency (i.e., the frequency of mutant allele) of 0.5. And the carrying capacity is established at 100, 000, with r = 1 and z = 1. In WF model, the population size at next generation is Nt+1 = Nt ×E(Kt). Note if the calculated value of Nt+1 is not an integer, we round it to the nearest whole number. The number of mutant alleles follows a binomial distribution, allowing the simulation of gene frequency overtime in WF model. For Haldane model, the number of offspring number is assumed to follow a beta-binomial distribution. The mean of offspring number E(Kt) is obtained from Eq. (5). The variance V(Kt) is obtained from Eq. (6) with parameters a = 1/300 and b = 1.2. Given the mean and variance of the beta-binomial distribution, we simulated the number of mutant alleles and population size overtime and then traced the gene frequency overtime.

Revised version: In both models, interactions between individuals are implicitly included through the dependency of the average number of offspring on population size, as defined by Eq. (5). This dependency leads to the logistic population growth, reflecting the density-dependent interactions. The initial population size is set to 4, with initial gene frequency (i.e., the frequency of mutant allele) of 0.5. And the carrying capacity is established at 100, 000, with r = 1 and z = 1. In WF model, the population size at next generation is Nt+1 = Nt ×E(Kt). Note if the calculated value of Nt+1 is not an integer, we round it to the nearest whole number. The number of mutant alleles follows a binomial distribution, allowing the simulation of gene frequency overtime in WF model. For Haldane model, the number of offspring number is assumed to follow a beta-binomial distribution. The mean of offspring number E(Kt) is obtained from Eq. (5). The variance V(Kt) is obtained from Eq. (6) with parameters a = 1/300 and b = 1.2. Given the mean and variance of the beta-binomial distribution, we simulated the number of mutant alleles and population size overtime and then traced the gene frequency overtime.

Fixation probability of a new mutation in Haldane model

Here we obtain the fixation probability of advantageous mutation in Haldane model governed by a branching process. The Haldane model considers a well-mixed population of Nt haploid parents with only two types of alleles (W for wildtype, M for mutant). There is no migration and new mutations in this model. Each individual is assumed to independently reproduce in discrete and non-overlapping generation, t = 0, 1, 2, …. Thus, the number of offspring per allele (also an individual in this haploid population) at any generation is represented by a set of identically and individually distributed random variables. In particular, the numbers of offspring of M and W alleles (denoted as KM and KW respectively) can be represented by following distributions.

For a particular distribution, we can obtain the average number of offspring as follows.

With the selection coefficient of s (0 for neutral mutation, positive for advantageous mutation, and negative for deleterious mutation) of M allele, average offspring number of M allele and W allele will follow the relationship.

Now, the evolution of the number of M allele (denoted as Mt) and the number of W allele (denoted as Wt) as time process is a branching process.

Earlier version: To compare Haldane model with WF model (the offspring number follows Poisson distribution with variance and mean equal to 1), we will set E(KW) = 1, E(KM) = 1 + s. And then let the offspring number follows a more general and realistic distribution (including negative binomial distribution (negBin), beta-binomial distribution (betaBin)), which can let the ratio of variance to mean range from 1 to a large number. (For simplicity, we let V(KM)/E(KM) = V(KW)/E(KW).) Based on the branching process, we obtained the fixation probability M alleles (see Supplementary Information for more details on the derivation). where

Revised version: To compare Haldane model with WF model (the offspring number follows Poisson distribution with variance and mean equal to 1), we will set E(KW) = 1, E(KM) = 1 + s. And then let the offspring number follows a more general and realistic distribution (including negative binomial distribution (negBin), beta-binomial distribution (betaBin)), which can let the ratio of variance to mean range from 1 to a large number. (For simplicity, we let V(KM)/E(KM) = V(KW)/E(KW).) Based on the branching process, we obtained the fixation probability M alleles (see Supplementary Information for more details on the derivation).

Added in the revised version: Where

Note uM(t = 1) = P(KM = 0) = i0, uW(t = 1) = P(KW = 0) = j0. And both {uM(t), t = 0, 1, 2, …} and {uW(t), t = 0, 1, 2, …} are a bounded monotonic sequence. Thus, although we cannot obtain the direct analytical solution for Pf, we could easily obtain their numerical solution by iteration to the case when both uM(t) and uW(t) converge.

Removed in the revised version: Supporting information